# Demo Code Challenge 

This demo is designed to give you the experience of a Code Challenge before your official Code Challenge. It provides practice with working through the Code Challenge, running tests to check your answers, and completing multiple item types.
 
*Read the instructions carefully.* Your code will need to meet detailed specifications to pass automated tests.

## Code Tests

We have provided some code tests for you to run to check that your work meets the item specifications. Passing these tests does not necessarily mean that you have gotten the item correct - there are additional hidden tests. However, if any of the tests do not pass, this tells you that your code is incorrect and needs changes to meet the specification. To determine what the issue is, read the comments in the code test cells, the error message you receive, and the item instructions.

## 1) Read `titanic.csv` into a pandas DataFrame named `df`

Use pandas to create a new DataFrame, called `df`, containing the data from the dataset in the file `titanic.csv` in the folder containing this notebook. 

Hint: Use the string `'./titanic.csv'` as the file reference.

```python
# Run this cell

import pandas as pd
import numpy as np
from numbers import Number
import warnings
warnings.filterwarnings('ignore')
```

**Starter Code**

```
df = 

```

<!-- grade_id: cell-f4647d39b17716ba, points: 0, grade: false, solution: true -->
```python
# your code here
df = pd.read_csv('./titanic.csv')
```

<!-- grade_id: cell-f7427f34ad14af7a, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a DataFrame named df

assert type(df) == pd.DataFrame

```

### 2) Create a variable `dogwood_index` containing the index for the `'dogwood'` element from the list `tree_type_names`

Below is the list `tree_type_names` that you will need - do not change this list.


```python
# Run this cell

tree_type_names = ['linden', 'spruce', 'dogwood', 'hickory', 'willow']
```

**Starter Code**

```
dogwood_index = 
```

<!-- grade_id: cell-0b84459c9aaab532, points: 0, grade: false, solution: true -->
```python
# your code here
dogwood_index = 2
```

<!-- grade_id: cell-00926387ac95d868, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named dogwood_index

assert isinstance(dogwood_index, Number)

```

## 3)  Briefly explain the purpose of code challenges in the Flatiron School Data Science program.

Enter a one sentence response below.

<!-- grade_id: cell-16aad1712359fe52, points: 0, grade: true, solution: true -->
To make sure student are ready for next phase and not getting lost.

```python

```