## Object Oriented Programming

In this exercise you will primarily be exploring a Pokemon dataset. Pokemon are fictional creatures from the [Nintendo franchise](https://en.wikipedia.org/wiki/Pok%C3%A9mon) of the same name.

Some Pokemon facts that might be useful:
* The word "pokemon" is both singular and plural. You may refer to "one pokemon" or "many pokemon".
* Pokemon have attributes such as a name, weight, and height.
* Pokemon have one or multiple "types". A type is something like "electric", "water", "ghost", or "normal" that indicates the abilities that pokemon may possess.
* The humans who collect pokemon are called "trainers".

As a pokemon trainer we want to make sure our pokemon are performing at their peak. To measure this, we want to calculate a pokemon's Body Mass Index (or BMI). This is a statistic calculated using the pokemon's height and weight. 

To help with this task we we will create Pokemon objects that methods can be called on. 

You'll be working with following dictionaries to create the `Pokemon` objects

```python
# Run this cell without changes
bulbasaur_data = {
    "name": 'bulbasaur',
    "weight": 69,
    "height": 7,
    "base_experience": 64,
    "types": ["grass", "poison"]
}
charmander_data = {
    "name": 'charmander',
    "weight": 85,
    "height": 6,
    "base_experience": 62,
    "types": ["fire"]
}
squirtle_data = {
    "name": 'squirtle',
    "weight": 90,
    "height": 5,
    "base_experience": 63,
    "types": ["water"]
}
```

### 1. Creating a Class

Create a class called `Pokemon` with an `__init__` method. 

Along with the necessary `self` parameter, the `__init__` method should 
take in `data` as a parameter.

With the idea that one of the dictionaries above will be passed in as `data`, 
assign these specific attributes within the `__init__` method:
 
* `name` : value from the 'name' key of the dictionary passed in `data`
* `weight`: value from the 'weight' key of the dictionary passed in `data`
* `height`: value from the 'height' key of the dictionary passed in `data`

<!-- grade_id: a3c7486e597da698c495a7e9019cf08b, points: 0, grade: false, solution: true -->
```python
# Create your class below with the correct syntax, including an __init__ method.
class Pokemon:
    def __init__(self, data):
        self.name = data['name']
        self.weight = data['weight']
        self.height = data['height']
# your code here
```

<!-- grade_id: eeed2f8de21f67ce0a7c2b163ce7430d, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# Pokemon should be a class that exists in this namespace
assert type(Pokemon) == type

```

    
### 2. Instantiating Objects

Using the `bulbasaur_data`, `charmander_data` and `squirtle_data` variables, create the corresponding pokemon objects.

<!-- grade_id: 48e5fe8a1075f95b1a9b46caedfbfc37, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code

bulbasaur = Pokemon(bulbasaur_data)
charmander = Pokemon(charmander_data)
squirtle = Pokemon(squirtle_data)

# your code here

# This code will test your implementation. Make sure the values printed
# match the dictionaries above!
def print_pokeinfo(pkmn):
    print('Name: ' + pkmn.name)
    print('Weight: ' + str(pkmn.weight))
    print('Height: ' + str(pkmn.height))
    print('\n')
    
print_pokeinfo(bulbasaur)
print_pokeinfo(charmander)
print_pokeinfo(squirtle)
```

Output:
```
Name: bulbasaur
Weight: 69
Height: 7


Name: charmander
Weight: 85
Height: 6


Name: squirtle
Weight: 90
Height: 5
```

<!-- grade_id: d858684ed54b6d901e6879ec5297a21b, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

assert type(bulbasaur) == Pokemon
assert type(charmander) == Pokemon
assert type(squirtle) == Pokemon

```

### 3. Update an instance attribute

Using the charmander instance, increase the weight attribute by 5. 

<!-- grade_id: cell-7e82657b2cb762a9, points: 0, grade: false, solution: true -->
```python
# your code here
charmander.weight += 5
charmander.weight
```

Output:
```
100
```

<!-- grade_id: cell-4f82e207aa09b98d, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

```

### 4. Instance Methods

Re-write the class `Pokemon` below, so that it has an instance method called `bmi` that calculates the BMI of a Pokemon.

BMI is defined by the formula: $\frac{weight}{height^{2}}$ 

The BMI should be calculated with weight in **kilograms** and height in **meters**. 

The height and weight data of Pokemon from the dictionaries above is in **decimeters** and **hectograms**, respectively. 

You will have to convert the given values of height and weight to kilograms and meters to make the BMI calculations correct.

For your convenience, here are the conversions:

```
1 decimeter = 0.1 meters
1 hectogram = 0.1 kilograms
```

**Don't forget**: since you are changing the `Pokemon` class, you will have to create **new objects** of this **new class**. 

If you use the objects created by the first class you wrote, they will not have the `bmi` instance!

You can assign these new objects the same names as you assigned the old ones:
`bulbasaur`, `charmander` and `squirtle`

<!-- grade_id: 2fdf7167c74891ba6f40289713cf8f1b, points: 0, grade: false, solution: true -->
```python
# Define the new Pokemon class here
class Pokemon:
    def __init__(self, data):
        self.name = data['name']
        self.weight = data['weight']
        self.height = data['height']
    
    def bmi(self):
        weight_in_kilograms = self.weight / 10
        height_in_meters = self.height / 10
        bmi = weight_in_kilograms / (height_in_meters**2)
        return bmi
            

# Replace None with appropriate code
bulbasaur = Pokemon(bulbasaur_data)
charmander = Pokemon(charmander_data)
squirtle = Pokemon(squirtle_data)

# your code here


# This code will test your implementation
print(bulbasaur.bmi()) 
print(charmander.bmi()) 
print(squirtle.bmi()) 
```

Output:
```
14.081632653061227
23.61111111111111
36.0
```

<!-- grade_id: 4bdd59541fbf9055ee7815c7e797521c, points: 2, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# bulbasaur.bmi should be a method on an instance of type Pokemon
import types
assert type(bulbasaur.bmi) == types.MethodType
# bulbasaur.bmi() should return a floating point number
assert type(bulbasaur.bmi()) == float
```

```python

```